# This is the main input for unit testing.
# Only this file is passed as command line argument.

rule_files:
  - ../rules/alerts.yml

evaluation_interval: 30s

tests:
  # Test 1.
  - interval: 30s
    # Series data. NB - the actual number is the total and must be divided by the number of threads
    input_series:
      - series: 'dmcproxy_process_cpu_usage_percent{job="redis", cluster="localhost", node="1"}'
        values: '70 80 90 90 90 90 90 90 90 90 90 90 90 80 70'

    # Unit test for alerting rules.
    alert_rule_test:
      # Unit test 1.
      - eval_time: 2m
        alertname: Proxy CPU Usage
        exp_alerts:
          # Alert 1.
          - exp_labels:
              severity: notification
              job: redis
              cluster: localhost
              redis: 1
            exp_annotations:
              summary: Shard CPU Usage
              description: "Busy Proxy - Cluster: localhost Node: 1 CPU: 90%"
              runbook: https://github.com/Redislabs-Solution-Architects/RedisEnterprisePrometheus/blob/master/runbooks/shards.md#shard-cpu-usage-is-excessive
