# This is the main input for unit testing.
# Only this file is passed as command line argument.

rule_files:
  - ../rules/alerts.yml

evaluation_interval: 30s

tests:
  # Test 1.
  - interval: 30s
    # Series data.
    input_series:
      - series: 'endpoint_read_requests{job="redis", cluster="localhost", db="1"}'
        values: '999999 2000000 2000001 2000001 2000001 2000001 2000001 2000001 2000001 2000001 2000001 2000001 2000000 2000000 999999'
      - series: 'endpoint_write_requests{job="redis", cluster="localhost", db="1"}'
        values: '999999 2000000 2000001 2000001 2000001 2000001 2000001 2000001 2000001 2000001 2000001 2000001 2000000 2000000 999999'

    # Unit test for alerting rules.
    alert_rule_test:
      # Unit test 1.
      - eval_time: 1m
        alertname: Excessive Redis Requests
        exp_alerts:
          # Alert 1.
          - exp_labels:
              severity: critical
              job: redis            
              cluster: localhost
              type: throughput
              db: 1
            exp_annotations:
              summary: Excessive Redis Requests
              description: "Too Many Redis Operations - Cluster: localhost DB: 1 2.000001e+06 (ops/sec)"
              runbook: https://redis.io/docs/latest/integrate/prometheus-with-redis-enterprise/observability/#connections
