# This is the main input for unit testing.
# Only this file is passed as command line argument.

rule_files:
  - ../rules/alerts.yml

evaluation_interval: 15s

tests:
  # Test 1.
  - interval: 15s
    # Series data.
    input_series:
      # - series: 'bdb_avg_latency{job="redis", cluster="localhost", db="1"}'
      - series: 'endpoint_acc_latency{job="redis", cluster="localhost", db="1"}'
        values: '0.05 0.06 0.05 0.06 0.07 0.08 0.09 0.08 0.07 0.06 0.05 0.06 0.07 0.06 0.05'
      - series: 'endpoint_total_started_res{job="redis", cluster="localhost", db="1"}'
        values: '0.005 0.006 0.005 0.006 0.007 0.008 0.009 0.008 0.007 0.006 0.005 0.006 0.007 0.006 0.005'

    # Unit test for alerting rules.
    alert_rule_test:
      # Unit test 1.
      - eval_time: 2m
        alertname: Average Latency Critical
        exp_alerts:
          # Alert 1.
          - exp_labels:
              severity: critical
              type: latency
              job: redis            
              cluster: localhost
              db: 1
            exp_annotations:
              summary: Average Latency Critical
              description: "High Latency - Cluster: localhost DB: 1 Latency: 10.000000000000002 ms"
              runbook: https://github.com/Redislabs-Solution-Architects/RedisEnterprisePrometheus/blob/master/runbooks/latency.md#average-latency
