# This is the main input for unit testing.
# Only this file is passed as command line argument.

rule_files:
  - ../rules/latency-alerts.yml

evaluation_interval: 15s

tests:
  # Test 1.
  - interval: 15s
    # Series data.
    input_series:
      # - series: 'bdb_avg_latency{job="redis", cluster="localhost", bdb="1"}'
      - series: 'endpoint_acc_latency{job="redis", cluster="localhost", db="1"}'
        values: '0.04 0.05 0.04 0.05 0.05 0.06 0.05 0.06 0.05 0.04 0.03 0.04 0.05 0.04 0.03'
      - series: 'endpoint_total_started_res{job="redis", cluster="localhost", db="1"}'
        values: '0.005 0.006 0.005 0.006 0.007 0.008 0.009 0.008 0.007 0.006 0.005 0.006 0.007 0.006 0.005'

    # Unit test for alerting rules.
    alert_rule_test:
      # Unit test 1.
      - eval_time: 2m
        alertname: Average Latency Warning
        exp_alerts:
          # Alert 1.
          - exp_labels:
              severity: notification
              job: redis            
              cluster: localhost
              type: latency
              db: 1
            exp_annotations:
              summary: Average Latency Warning
              description: "High Latency - Cluster: localhost DB: 1 Latency: 7.142857142857143 ms"
              runbook: https://redis.io/docs/latest/integrate/prometheus-with-redis-enterprise/observability/#read-latency

