# This is the main input for unit testing.
# Only this file is passed as command line argument.

rule_files:
  - ../rules/utilization-alerts.yml

evaluation_interval: 1m

tests:
  # Test 1.
  - interval: 1m
    # Series data.
    input_series:
      - series: 'redis_server_keyspace_read_hits{job="redis", role="master", cluster="localhost", db="1"}'
        values: '1 2 3 2 3 2 3 4 3 2 3 2 1 1 3'
      - series: 'redis_server_keyspace_read_misses{job="redis", role="master", cluster="localhost", db="1"}'
        values: '10 20 10 20 10 10 10 20 20 20 20 10 10 10 5'

    # Unit test for alerting rules.
    alert_rule_test:
      # Unit test 1.
      - eval_time: 1m
        alertname: Low Hit Ratio
        exp_alerts:
          # Alert 1.
          - exp_labels:
              severity: notification
              job: redis            
              role: master
              type: utilization
              cluster: localhost
              db: 1
            exp_annotations:
              summary: Low Hit Ratio
              description: "Low Hit Ratio: localhost DB: 1 Hit Ratio: 0.09090909090909091%"
              runbook: https://redis.io/docs/latest/integrate/prometheus-with-redis-enterprise/observability/#cache-hit-ratio-and-eviction
