# This is the main input for unit testing.
# Only this file is passed as command line argument.

rule_files:
  - ../rules/shard-alerts.yml

evaluation_interval: 1m

tests:
  # Test 1.
  - interval: 1m
    # Series data.
    input_series:
      - series: 'redis_server_master_link_status{role="slave", cluster="localhost", job="redis", node="1", redis="1"}'
        values: '0 0 0 0 0 0 0 0 0 0 0 0 0 0 0 0 0 0 0 0'

    # Unit test for alerting rules.
    alert_rule_test:
      # Unit test 1.
      - eval_time: 1m
        alertname: Master Shard Down
        exp_alerts:
          # Alert 1.
          - exp_labels:
              severity: notification
              job: redis            
              cluster: localhost
              role: slave
              type: shard
              node: 1
              redis: 1
            exp_annotations:
              summary: Master Shard Down
              description: "Slave has no master - Cluster: localhost Shard: 1 Node: 1"
              runbook: https://redis.io/docs/latest/integrate/prometheus-with-redis-enterprise/observability/#cpu

