# This is the main input for unit testing.
# Only this file is passed as command line argument.

rule_files:
  - ../rules/alerts.yml

evaluation_interval: 1m

tests:
  # Test 1.
  - interval: 1m
    # Series data.
    input_series:
      - series: 'namedprocess_namegroup_cpu_seconds_total{job="redis", cluster="localhost", node="1"}'
        # values: '70000 80000 90000 100000 110000 120000 130000 140000 150000 160000 170000 180000 190000 200000 210000'
        ### 20% increase/minute
        # values: '70000 84000 1008000 120960 145152 174182 209018 250822 300987 361184 433421 520106 624127 748952'
        ### 10% increase/minute
        values: '50000 55000 60500 66550 73205 80525 88578 97436 107180 117897 129687 142656 156920 172613'
        ### 5% increase/minute
        # values: '50000 52500 55125 57881 60775 63814 67004 70354 73872 77566 81444 85516 89792 94282'

#    promql_expr_test:
#      - expr: rate(namedprocess_namegroup_cpu_seconds_total{job="redis", cluster="localhost", node="1"}[1m])
#        eval_time: 10m
#        exp_samples:
#          - value: 10000

    # Unit test for alerting rules.
    alert_rule_test:
      # Unit test 1.
      - eval_time: 10m
        alertname: Node CPU Usage
        exp_alerts:
          # Alert 1.
          - exp_labels:
              severity: critical
              cluster: localhost
              job: redis
              type: node
              node: 1
            exp_annotations:
              summary: Node CPU Usage
              description: "High CPU Usage - Cluster: localhost Usage: 171.2875%"
              runbook: https://github.com/Redislabs-Solution-Architects/RedisEnterprisePrometheus/blob/master/runbooks/nodes.md#node-cpu-usage
